# Assignment 2

For this assignment you'll be looking at 2017 data on immunizations from the CDC. Your datafile for this assignment is in [assets/NISPUF17.csv](assets/NISPUF17.csv). A data users guide for this, which you'll need to map the variables in the data to the questions being asked, is available at [assets/NIS-PUF17-DUG.pdf](assets/NIS-PUF17-DUG.pdf). **Note: you may have to go to your Jupyter tree (click on the Coursera image) and navigate to the assignment 2 assets folder to see this PDF file).**

## Question 1
Write a function called `proportion_of_education` which returns the proportion of children in the dataset who had a mother with the education levels equal to less than high school (<12), high school (12), more than high school but not a college graduate (>12) and college degree.

*This function should return a dictionary in the form of (use the correct numbers, do not round numbers):* 
```
    {"less than high school":0.2,
    "high school":0.4,
    "more than high school but not college":0.2,
    "college":0.2}
```

<!-- grade_id: cell-eea16d020eb52ae7, points: 0, grade: false, solution: true -->
```python
import pandas as pd
import numpy as np

#Loading ds
df = pd.read_csv('./assets/NISPUF17.csv')

values = df['EDUC1'].value_counts().sort_index()

#calculations
ratios = values.values / sum(values.values)
print('ratios: ', ratios)

keys = ['less than high school', 'high school', 'more than high school but not college', 'college']

#dictionary for mapping keys to ratios
res = {key: ratio for key, ratio in zip(keys, ratios)}
print(res)
print(type(res))
print(res.keys())

def proportion_of_education():
    keys = ['less than high school', 'high school', 'more than high school but not college', 'college']
    res = {key: ratio for key, ratio in zip(keys, ratios)}
    return res

```

Output:
```
ratios:  [0.10202002 0.17235201 0.24588091 0.47974706]
{'less than high school': 0.10202002459160373, 'high school': 0.172352011241876, 'more than high school but not college': 0.24588090637625154, 'college': 0.47974705779026877}
<class 'dict'>
dict_keys(['less than high school', 'high school', 'more than high school but not college', 'college'])
```

<!-- grade_id: cell-c0eeef201366f51c, points: 1, grade: true, solution: false -->
```python
assert type(proportion_of_education())==type({}), "You must return a dictionary."
assert len(proportion_of_education()) == 4, "You have not returned a dictionary with four items in it."
assert "less than high school" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "high school" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "more than high school but not college" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "college" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."

```

## Question 2

Let's explore the relationship between being fed breastmilk as a child and getting a seasonal influenza vaccine from a healthcare provider. Return a tuple of the average number of influenza vaccines for those children we know received breastmilk as a child and those who know did not.

*This function should return a tuple in the form (use the correct numbers:*
```
(2.5, 0.1)
```

<!-- grade_id: cell-77f18c512324eabb, points: 0, grade: false, solution: true -->
```python
import pandas as pd
import numpy as np

# Load data:
df = pd.read_csv('./assets/NISPUF17.csv')
#print(df.head())

#breastfeeding == yes - page 52
#print(df['CBF_01'])
df_bfy = df[df['CBF_01'] == 1]
print('Amount of breastfed kids:', len(df_bfy))

#breastfeeding == no
df_bfn = df[df['CBF_01'] == 2]
print('Amount of kids not breastfed:', len(df_bfn))

#'P_NUMFLU' total number of seasonal influenza doses + fill missing values - replace with 0
df_inf = df['P_NUMFLU']
print(df_inf)
df['P_NUMFLU'] = df['P_NUMFLU'].fillna(0)
print(df['P_NUMFLU'])

#separate for 2 groups based brestfed yes/no
df_bf_vac = df_bfy['P_NUMFLU']
df_nbf_vac = df_bfn['P_NUMFLU']

#Calculations for each group
avg_bf_vac = df_bf_vac.mean()
avg_nbf_vac = df_nbf_vac.mean()

#final conclusion and result 
print('Average number of influenza vaccines for breastfed kids:', avg_bf_vac)
print('Average number of influenza vaccines for non breastfed kids:', avg_nbf_vac)

def average_influenza_doses():
    bf = avg_bf_vac
    nbf = avg_nbf_vac
    res = (bf, nbf)
    return res
    raise NotImplementedError()
```

Output:
```
Amount of breastfed kids: 24261
Amount of kids not breastfed: 4085
0        NaN
1        NaN
2        NaN
3        3.0
4        0.0
        ... 
28460    NaN
28461    NaN
28462    NaN
28463    NaN
28464    NaN
Name: P_NUMFLU, Length: 28465, dtype: float64
0        0.0
1        0.0
2        0.0
3        3.0
4        0.0
        ... 
28460    0.0
28461    0.0
28462    0.0
28463    0.0
28464    0.0
Name: P_NUMFLU, Length: 28465, dtype: float64
Average number of influenza vaccines for breastfed kids: 1.8799187420058687
Average number of influenza vaccines for non breastfed kids: 1.5963945918878317
```

<!-- grade_id: cell-54a3ba6cff31caa7, points: 1, grade: true, solution: false -->
```python
assert len(average_influenza_doses())==2, "Return two values in a tuple, the first for yes and the second for no."

```

## Question 3
It would be interesting to see if there is any evidence of a link between vaccine effectiveness and sex of the child. Calculate the ratio of the number of children who contracted chickenpox but were vaccinated against it (at least one varicella dose) versus those who were vaccinated but did not contract chicken pox. Return results by sex. 

*This function should return a dictionary in the form of (use the correct numbers):* 
```
    {"male":0.2,
    "female":0.4}
```

Note: To aid in verification, the `chickenpox_by_sex()['female']` value the autograder is looking for starts with the digits `0.0077`.

<!-- grade_id: cell-a0a9e6fe67698006, points: 0, grade: false, solution: true -->
```python
df = pd.read_csv('./assets/NISPUF17.csv')
print('Number of total records: ', len(df))

#Received at least one dose:
vaccinated = df[df['P_NUMVRC'] >= 1]
print(vaccinated.head())
print(len(vaccinated))

#number of children who contracted chickenpox CPOX Chicken Pox Variables - page 52
contracted_chickenpox = vaccinated[vaccinated['HAD_CPOX'] == 1]
print(contracted_chickenpox.head())
print('Number of children who contracted chikcenpox:')
print(len(contracted_chickenpox))

#children who were vaccinated but did not contract chickenpox 
not_contracted_chickenpox = vaccinated[vaccinated['HAD_CPOX'] == 2]
print(not_contracted_chickenpox.head())
print('number of children who did not contract chickenpox:', len( not_contracted_chickenpox))

def chickenpox_by_sex():
    #ratios by sex
    ratio_male = len(contracted_chickenpox[contracted_chickenpox['SEX'] == 1]) / len(not_contracted_chickenpox[not_contracted_chickenpox['SEX'] == 1])
    ratio_female = len(contracted_chickenpox[contracted_chickenpox['SEX'] == 2]) / len(not_contracted_chickenpox[not_contracted_chickenpox['SEX'] == 2])
    print(ratio_male)
    print(ratio_female)
    res ={"male": ratio_male, "female": ratio_female}
    return res
    raise NotImplementedError()

#result = chickenpox_by_sex()
#print(result)
```

Output:
```
Number of total records:  28465
    Unnamed: 0  SEQNUMC  SEQNUMHH  PDAT    PROVWT_D     RDDWT_D  STRATUM  \
3            4    86131      8613     1  675.430817  333.447418     1040   
6            7   221381     22138     1  516.427723  248.916038     1008   
10          11   128211     12821     1  335.398714  198.123294     1057   
12          13   137041     13704     1   82.582116   45.544717     1051   
16          17   188611     18861     1  345.352030  196.111859     1057   

    YEAR  AGECPOXR  HAD_CPOX  ...  XVRCTY2  XVRCTY3  XVRCTY4  XVRCTY5  \
3   2017       NaN         2  ...                        NaN      NaN   
6   2017       NaN         2  ...                        NaN      NaN   
10  2017       NaN         2  ...                        NaN      NaN   
12  2017       NaN         2  ...                        NaN      NaN   
16  2017       NaN         2  ...       VO               NaN      NaN   

    XVRCTY6  XVRCTY7  XVRCTY8  XVRCTY9  INS_STAT2_I  INS_BREAK_I  
3    
```

<!-- grade_id: cell-c4f1714db100c865, points: 1, grade: true, solution: false -->
```python
assert len(chickenpox_by_sex())==2, "Return a dictionary with two items, the first for males and the second for females."

```

## Question 4
A correlation is a statistical relationship between two variables. If we wanted to know if vaccines work, we might look at the correlation between the use of the vaccine and whether it results in prevention of the infection or disease [1]. In this question, you are to see if there is a correlation between having had the chicken pox and the number of chickenpox vaccine doses given (varicella).

Some notes on interpreting the answer. The `had_chickenpox_column` is either `1` (for yes) or `2` (for no), and the `num_chickenpox_vaccine_column` is the number of doses a child has been given of the varicella vaccine. A positive correlation (e.g., `corr > 0`) means that an increase in `had_chickenpox_column` (which means more no’s) would also increase the values of `num_chickenpox_vaccine_column` (which means more doses of vaccine). If there is a negative correlation (e.g., `corr < 0`), it indicates that having had chickenpox is related to an increase in the number of vaccine doses.

Also, `pval` is the probability that we observe a correlation between `had_chickenpox_column` and `num_chickenpox_vaccine_column` which is greater than or equal to a particular value occurred by chance. A small `pval` means that the observed correlation is highly unlikely to occur by chance. In this case, `pval` should be very small (will end in `e-18` indicating a very small number).

[1] This isn’t really the full picture, since we are not looking at when the dose was given. It’s possible that children had chickenpox and then their parents went to get them the vaccine. Does this dataset have the data we would need to investigate the timing of the dose?

<!-- grade_id: cell-8afff07f564cf79a, points: 0, grade: false, solution: true -->
```python
#From page 52:
#HAD_CPOX - did child ever have chicken pox
#P_NUMVRC - total number of varicella doses
df = pd.read_csv('./assets/NISPUF17.csv')
print(df.head())

df_filtered = df[['HAD_CPOX', 'P_NUMVRC']].dropna()
print(df_filtered.head())
print('ALl values of HAD_CPOX variable:', df_filtered['HAD_CPOX'].unique())

#cleaning 
df_valid_filtered = df_filtered[(df_filtered['HAD_CPOX'] == 1) | (df_filtered['HAD_CPOX'] == 2)]
print(df_valid_filtered['HAD_CPOX'].unique())

print('ALl values of P_NUMVRC variable:', df_filtered['P_NUMVRC'].unique())


def corr_chickenpox():
    import scipy.stats as stats
    corr, pval=stats.pearsonr(df_valid_filtered["HAD_CPOX"],df_valid_filtered["P_NUMVRC"])
    return corr

    raise NotImplementedError()
```

Output:
```
Unnamed: 0  SEQNUMC  SEQNUMHH  PDAT    PROVWT_D     RDDWT_D  STRATUM  YEAR  \
0           1   128521     12852     2         NaN  235.916956     1031  2017   
1           2    10741      1074     2         NaN  957.353840     1068  2017   
2           3   220011     22001     2         NaN  189.611299     1050  2017   
3           4    86131      8613     1  675.430817  333.447418     1040  2017   
4           5   227141     22714     1  482.617748  278.768063     1008  2017   

   AGECPOXR  HAD_CPOX  ...  XVRCTY2  XVRCTY3  XVRCTY4  XVRCTY5  XVRCTY6  \
0       NaN         2  ...                        NaN      NaN      NaN   
1       NaN         2  ...                        NaN      NaN      NaN   
2       NaN         2  ...                        NaN      NaN      NaN   
3       NaN         2  ...                        NaN      NaN      NaN   
4       NaN         2  ...                        NaN      NaN      NaN   

   XVRCTY7  XVRCTY8  XVRCTY9  INS_STAT2_I  INS_BREAK_I  
0      N
```

<!-- grade_id: cell-73408733533a29a5, points: 1, grade: true, solution: false -->
```python
assert -1<=corr_chickenpox()<=1, "You must return a float number between -1.0 and 1.0."

```